---
title: 'Test chapitre'
description: 'Chapter description goes here.'
free_preview: true
---

## Video

```yaml
type: VideoExercise
key: b03597fe1e
xp: 50
```

`@projector_key`
0239f569a5f24fe2e17551364b65f3ff

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 91559b3612
xp: 100
```

So this is an exercise or something, i don't really know

`@instructions`
instructionssssssssssssssssssssssssss

`@hint`
a hintt t t t t t t t tt t tt . t tt t t tt

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# print the number 5 here


```

`@solution`
```{python}
print(5)
```

`@sct`
```{python}
Ex().has_output('5')
success_msg('well done!')
```
